<a href="https://colab.research.google.com/github/suchirasuresh/Applied-Data-Science-in-Medicine-Psychology-ads_assignment_1--suchira_suresh-/blob/main/04_python_basics_III.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Fill in your github username and idm credentials below:

```python
github_username = "suchirasuresh"
idm_credentials = "yh07ajaw"
```

---

# Applied Data Science in Medicine & Psychology - Assignment 4
## Python Basics III
---
**Deadline: 23.05.2023, 23:59**
### **General setup and submission**
1. Save an editable copy of this notebook to your Google Drive ``File`` $\rightarrow$ ``Save a copy in Drive`` and remove the automatically appended "Copy of" text. **Do not change the file name**.
2. Work on the assignment in the *copied* version of the notebook. Make sure that you **replace** all parts of the tasks indicated with *YOUR CODE HERE* and **raise NotImplementedError()** (otherwise this causes an error in your implementation) or *YOUR ANSWER HERE* with your solution.
3. If you take a break during the implementation, you can save the current status of your work in your repository following the submission instructions below. When you continue working on the assignment, do this in the version of the notebook that you saved in your Google Drive (access from your Google Colab account) or your repository.
4. After finishing the assignment submit your assignment as follows:
    * Check that everything runs as expected. To do that select in the menubar ``Runtime`` $\rightarrow$ ``Restart and run all``. This will clear all your local variables and runs your notebook from the beginning
    * Save your notebook, click ``File`` $\rightarrow$ ``Save``
    * Click ``File`` $\rightarrow$ ``Save a copy in GitHub``
    * Select ``digikolleg-data-science-psychology/04_python_basics_III-<username>`` from the Repository drop-down menu
    * Type ``04_python_basics_III.ipynb`` into the File path text box
    * Optionally, you may enter a description into the "Commit message" text box (like: "solved task 1" or "ready for grading")
    * Make sure that the "Include a link to Colaboratory" box is checked
    * Press "Ok"
    * A new tab or window should open up, displaying your submitted notebook within your GitHub repository

Before the assignment deadline, you can re-submit your assignment as often as you would like. This will update the copy of your assignment file in your assignment's repository.

### **Assignment Description**
In the third Python basics assignment, you will learn about the advanced usage of loops and list comprehension to improve code efficiency and readability. Also, you will deepen your understanding of the order of operations and practice error handling techniques, such as try-except blocks and raising exceptions, to identify and deal with errors in Python code. 

#### **Assignment goals**
You are going to learn the following things in this assignment:
1. how to improve the efficiency of your code with the advanced usage of loops
2. how to create lists using a single line of code
3. deepen your understanding of the order of operations
4. identifying and dealing with errors

#### **Assignment point value**
This assignment is worth a total (up to) **10** points. If your code passes all of the manual and automatic tests you will receive full credit.

### **Tasks**
Make sure that you now work on your *copied* version of this assignment's notebook.

#### 0. Imports and Setup
Please **run** the cell below to import the libraries needed for this assignment.

```python
import numpy as np
from warnings import warn, catch_warnings
from contextlib import suppress
```

#### 1. Loops
Python has some built-in functions/statements like ``range()``, ``continue``, and ``pass`` which are useful tools for controlling the flow of loops and performing specific actions based on certain conditions. 

##### 1.1 Play around with the ``range()`` function:
* Write a for loop that prints every third number between 0 and 20:

<!-- grade_id: range_1, points: 0, grade: false, solution: true -->
```python
print(list(range(3, 20, 3)))
```

Output:
```
[3, 6, 9, 12, 15, 18]
```

* Write a for loop that prints every second number between -5 and 5 (inclusive endpoints) in reverse order:

<!-- grade_id: range_2, points: 0, grade: false, solution: true -->
```python
for i in range (5, -6, -2):
  print((i))

```

Output:
```
5
3
1
-1
-3
-5
```

* Create a list ``my_list`` with every 4th number between -10 and 10 (inclusive endpoints):

<!-- grade_id: range_3, points: 0, grade: false, solution: true -->
```python

my_list= list(range(-10, 11, 4))
print(my_list)

  

```

Output:
```
[-10, -6, -2, 2, 6, 10]
```

<!-- grade_id: range_3_test, points: 0.5, grade: true, solution: false -->
```python
assert len(my_list) == 6, "my_list must be of length 6. Check if you handled your endpoints correctly"
assert my_list[0] == -10 and my_list[5] == 10, "Check if you handled your endpoints correctly"
```

##### 1.2 Play around with ``pass``, ``continue``, and ``break`` statements:

``break``: The break statement is used to exit a loop prematurely. When encountered, the break statement causes the loop to immediately terminate and the program execution to continue with the next statement following the loop.

* What happens if you replace ``pass`` with ``continue`` or ``break``in the following code cell?

```python
for i in range(10):
    if i == 6:
      break
    print(i) 
print("executed loop")
```

Output:
```
0
1
2
3
4
5
executed loop
```

* Can you explain the difference between these 3 statements?
1.pass keyword is an empty instruction and the interpreter executes the next instructions after this
2.break terminates the loop. Hence output is 0 to 5 followed by printing "executed loop"
2.Continue keyword continues the loop if the above instruction to it is satisfied. In the above example we get same output for continue and pass.However, thats not the same everywhere.


*YOUR ANSWER HERE*

#### 2. List Comprehension
List comprehension is a concise way to create a new list by applying a transformation or filtering operation to an existing list or another iterable object. It has a compact syntax that allows you to create a new list in a single line of code.

Here is the basic syntax for list comprehension in Python:

```
y = [my_fun(i) for i in x if condition(i)]
```

* ``y`` is the newly created list
* ``x`` is a list of values that y will be based on
* ``i`` refers to each individual element of x
* ``my_fun`` is a function that is applied to each element of x
* ``condition`` is an optional statement that modifies which elements of x are included (only elements i where condition(i) is True will be used)

##### 2.1 Create a list ``squares`` that contains the squared values of 1 to 10 (inclusive endpoints) using list comprehension:

<!-- grade_id: list_com_1, points: 0, grade: false, solution: true -->
```python
squares= [i*i for i in range(1,11) ]
print (squares)
```

Output:
```
[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

<!-- grade_id: list_com_1_test, points: 0.5, grade: true, solution: false -->
```python
assert len(squares) == 10, "squares must be of length 10. Check if you handled your endpoints correctly"
assert squares[9] == 100, "Check if you handled your endpoints correctly"
```

##### 2.2 Create a list ``even_numbers`` which contains the even numbers between 1 and 10 (inclusive endpoints if applicable) using list comprehension:

<!-- grade_id: list_com_2, points: 0, grade: false, solution: true -->
```python
even_numbers= [i for i in range(1,11) if i%2==0]
print (even_numbers)
```

Output:
```
[2, 4, 6, 8, 10]
```

<!-- grade_id: list_com_2_test, points: 0.5, grade: true, solution: false -->
```python
assert len(even_numbers) == 5, "even_numbers must be of length 5. Check if you handled your endpoints correctly"
assert even_numbers[0] == 2 and even_numbers[4] == 10, "Check if you handled your endpoints correctly"
```

##### 2.3 Create a list ``ab_sum`` which contains the elementwise sum of the given lists ``a`` and ``b`` using list comprehension:
**Hint**: Look into ``len(my_list)`` which returns the number of elements of the given ``my_list``.

<!-- grade_id: list_com_3, points: 0, grade: false, solution: true -->
```python
a = [1, 3, 5, 7, 9]
b = [2, 4, 6, 8, 10]
ab_sum= [a[i]+b[i] for i in range(len(a))]
print(ab_sum)
```

Output:
```
[3, 7, 11, 15, 19]
```

<!-- grade_id: list_com_3_test, points: 0.5, grade: true, solution: false -->
```python
assert len(ab_sum) == 5, "ab_sum must be of length 5. Check if you handled your endpoints correctly"
assert ab_sum[0] == 3 and ab_sum[4] == 19, "Check if you handled your endpoints correctly"
```

#### 3. Study Name Conformity Check
In the following, you will implement 3 different functions which later on will be used to perform a conformity check of a random study name. The requirements for a study name in an artificial database are:
* req_1: only lowercase letters are allowed
* req_2: only letters, numbers, and underscore \'_\' are allowed
* req_3: the study name consists of 10 to 50 characters

##### 3.1 Create a function ``check_lower`` with ``name`` as parameter which checks whether all letters are lowercase (req_1):
* if **all** letters are lowercase return ``True``
* if **any** letter is a captial letter return ``False``

**Hint**: Check out if one of the built-in [string methods](https://docs.python.org/3/library/stdtypes.html#string-methods) can help you with the implementation.

<!-- grade_id: check_lower, points: 0, grade: false, solution: true -->
```python
def check_lower(name):
  return name.islower()

```

<!-- grade_id: check_lower_test, points: 1, grade: true, solution: false -->
```python
assert check_lower("study_psyc_1") == True, "Check if you return the correct boolean value"
assert check_lower("Study_Psyc_2") == False, "Check if you return the correct boolean value and used the right function"
```

##### 3.2 Create a function ``check_special_characters`` with ``name`` as parameter which checks whether all elements are letters, numbers, or underscore \'_\' (req_2):
* if **all** letters meet the requirement return ``False``
* if **any** letter does not meet the requirement return ``True``

**Hint**: Check again the [string methods](https://docs.python.org/3/library/stdtypes.html#string-methods) for a useful function.

<!-- grade_id: check_special_character, points: 0, grade: false, solution: true -->
```python
def check_special_characters(name):
  for c in name:
    if not (c.isalnum() and c!='_'):
      output=True
    else:
      output=False
  return output

```

<!-- grade_id: check_special_character_test, points: 1, grade: true, solution: false -->
```python
assert check_special_characters("study_1") == False, "Check if you return the correct boolean value"
assert check_special_characters("study_*") == True, "Check if you return the correct boolean value and handled special characters correctly"
assert check_special_characters("*study_") == True, "Check if you return the correct boolean value and handled special characters correctly"
```

##### 3.3 Create a function ``check_length`` with ``name`` as parameter which checks whether the name has at least 10 but maximum 50 characters (req_3):
* if name consist of 10 to 50 characters (inclusive endpoints) return ``True``
* if name has more or less characters return ``False``

<!-- grade_id: check_length, points: 0, grade: false, solution: true -->
```python
def check_length(name):
  if (len(name)>=10 and len(name)<=50):
    output =True
  else:
    output= False
  return output


```

<!-- grade_id: check_length_test, points: 1, grade: true, solution: false -->
```python
assert check_length("psychology") == True, "Check whether you handled your endpoints correctly"
assert check_length("study_applied_data_science_in_medicine_and_psychology") == False, "Check if you return the correct boolean value"
```

##### 3.4 Create a function ``check_study_name`` with ``name`` as parameter and perform the study name conformity check using the functions implemented above:
* Check if ``name`` is of type string and meets req_1, req_2, and req_3 using the functions above
* if ``name`` meets all requirements return ``True``
* if ``name`` fails one of the tests return ``False``
* Can you implement this in only one line of code?

**Hint**: check out the [isinstance()](https://docs.python.org/3/library/functions.html#isinstance) function to determine the type of parameters.

<!-- grade_id: check_study_name, points: 0, grade: false, solution: true -->
```python

def check_study_name(name):
  x= isinstance(name,str)

  output= (x and check_lower(name)) and (check_special_characters(name) and check_length(name))
  return output
  """
  if not (((x==False) or check_lower(name) or check_special_characters(name) or check_length(name))) is True:
    output = False
  else:
    output = True

    return output
    """



```

<!-- grade_id: check_study_name_test, points: 1, grade: true, solution: false -->
```python
assert check_study_name("study_psyc_2") == True, "Ensure that you combined the functions correctly using logical operators"
assert check_study_name(1) == False, "Ensure that you checked the data type of name"
assert check_study_name("Study_1") == False, "Ensure that you performed the check for lowercase letters"
assert check_study_name("study_psyc_#1") == False, "Ensure that you performed special characters check"
assert check_study_name("study_2") == False, "Ensure that you performed the length check"
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-136-b8ce419fadfc> in <cell line: 1>()
----> 1 assert check_study_name("study_psyc_2") == True, "Ensure that you combined the functions correctly using logical operators"
      2 assert check_study_name(1) == False, "Ensure that you checked the data type of name"
      3 assert check_study_name("Study_1") == False, "Ensure that you performed the check for lowercase letters"
      4 assert check_study_name("study_psyc_#1") == False, "Ensure that you performed special characters check"
      5 assert check_study_name("study_2") == False, "Ensure that you performed the length check"

AssertionError: Ensure that you combined the functions correctly using logical operators
```

#### 4. Error Handling
When a program encounters an error, it provides valuable feedback about what went wrong and where the problem occurred. This information can be used to diagnose and fix bugs in the code. Error handling also allows you to handle unexpected situations and lets your program continue to function even in the face of errors and other unexpected events.

##### 4.1 In this task you will implement functions for a "safe" division:
* Create a function ``safe_division`` with ``dividend`` and ``divisor`` as parameters. Return the quotient of these two. Then investigate which Exceptions might occur during the division using the test cases presented below. Implement your error handling accordingly such that ``np.nan`` (floating point representation of Not a Number (NaN)) is returned when an Exception occurs. You can use the ``finally`` block for example to print the result of your safe division:

**Hint**: Check out the lecture slides for a suitable Python construct.

<!-- grade_id: safe_division, points: 0, grade: false, solution: true -->
```python
def safe_division(dividend, divisor):
    try:
        quotient = dividend / divisor
    except Exception:
        quotient = float('nan')
    except ZeroDivisionError:
        quotient = float('nan')
    finally:
        print("Result of safe division:", quotient)
        return quotient

```

<!-- grade_id: safe_division_test, points: 1, grade: true, solution: false -->
```python
assert safe_division(10, 1) == 10, "Ensure that you perform the division and return the quotient"
assert np.isnan(safe_division(10, "1")), "Ensure that you handle this Exception correctly"
assert np.isnan(safe_division(10, 0)), "Ensure that you handle this Exception correctly"
```

Output:
```
Result of safe division: 10.0
Result of safe division: nan
Result of safe division: nan
```

* Now create a function ``division`` with ``dividend`` and ``divisor`` as parameters.
    * Check the datatype of ``dividend``. If it is not of type int or float print a meaningful message and return
    * Check the datatype of ``divisor``. If it is not of type list typecast the parameter to list
    * Then divide ``dividend`` with every element of divisor using ``safe_divison``
    * Return a list of all the quotients
    
**Hint**: Try ``[...]`` instead of ``list()`` for the typecast of ``divisor``.

<!-- grade_id: division, points: 0, grade: false, solution: true -->
```python
def division(dividend, divisor):
    if not isinstance(dividend, (int, float)):
        print("type int or float allowed")
        return
    
    if not isinstance(divisor, list):
        divisor = list(divisor)
    
    final_quotients = []
    for c in divisor:
        quotient = safe_division(dividend, c)
        final_quotients.append(quotient)
    
    return final_quotients

```

<!-- grade_id: division_test, points: 1, grade: true, solution: false -->
```python
assert division("1", 1) == None, "Ensure that you performed the datatype check of dividend"
assert division(1, 1) == [1.0], "Ensure that you performed the datatype check of divisor"
assert division(10, [2, "hi", 2.5, 0, [1, 2], 200, None]) == [5.0, np.nan, 4.0, np.nan, np.nan, 0.05, np.nan], "Check if you performed the division elementwise"
assert division(0, [0, 1]) == [np.nan, 0], "Check if you performed the division elementwise"
```

Output:
```
type int or float allowed
<ipython-input-153-232e41ff78b4> in <cell line: 2>()
      1 assert division("1", 1) == None, "Ensure that you performed the datatype check of dividend"
----> 2 assert division(1, 1) == [1.0], "Ensure that you performed the datatype check of divisor"
      3 assert division(10, [2, "hi", 2.5, 0, [1, 2], 200, None]) == [5.0, np.nan, 4.0, np.nan, np.nan, 0.05, np.nan], "Check if you performed the division elementwise"
      4 assert division(0, [0, 1]) == [np.nan, 0], "Check if you performed the division elementwise"

<ipython-input-152-33f03c4811ac> in division(dividend, divisor)
      5 
      6     if not isinstance(divisor, list):
----> 7         divisor = list(divisor)
      8 
      9     final_quotients = []

TypeError: 'int' object is not iterable
```

##### 4.2 Create a function ``check_study`` with the 2 parameters ``name`` and ``number_participants`` which checks if the study parameters are valid:
* ``name``
    * if it is ``None`` raise a warning with a meaningful message
    * otherwise, use the ``check_study_name()`` function you created before to test for study name requirements. If they are not met, raise the Exception ``ValueError`` with a meaningful message
* ``number_participants``
    * if it is not of type ``int`` raise the Exception ``TypeError`` with a meaningful message
    * if it is negative raise the Exception ``ValueError`` with a meaningful message

<!-- grade_id: check_study, points: 0, grade: false, solution: true -->
```python
def check_study(name, number_participants):
  if name is None:
    warn("enter Name")
  elif check_study_name(name):
    raise ValueError("Name Invalid")

  if not isinstance(number_participants, int):
    TypeError("number of participants mustbe integer")

  if number_participants<0:
    raise ValueError("number of participants mustbe positive")
```

<!-- grade_id: check_study_test, points: 2, grade: true, solution: false -->
```python
with catch_warnings(record=True) as w: check_study(None, 2); assert len(w) == 1, "Check if you used a Warning"
with suppress(ValueError): check_study("study_psyc_#2", 2); assert False, "Check if you raise the correct Exception for study name"
with suppress(TypeError): check_study("study_psyc_1", [12]); assert False, "Check if you raise the correct Exception for number_participants"
with suppress(ValueError): check_study("study_psyc_1", -1); assert False, "Check if you raise the correct Exception for number_participants"
```