<a href="https://colab.research.google.com/github/suchirasuresh/Applied-Data-Science-in-Medicine-Psychology-ads_assignment_1--suchira_suresh-/blob/main/05_dicts_classes.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Fill in your github username and idm credentials below:

```python
github_username = "suchirasuresh"
idm_credentials = "22770770"
```

---

# Applied Data Science in Medicine & Psychology - Assignment 5
## Dictionaries and Classes
---
**Deadline: 06.06.2023, 23:59**
### **General setup and submission**
1. Save an editable copy of this notebook to your Google Drive ``File`` $\rightarrow$ ``Save a copy in Drive`` and remove the automatically appended "Copy of" text. **Do not change the file name**.
2. Work on the assignment in the *copied* version of the notebook. Make sure that you **replace** all parts of the tasks indicated with *YOUR CODE HERE* and **raise NotImplementedError()** (otherwise this causes an error in your implementation) or *YOUR ANSWER HERE* with your solution.
3. If you do a break during the implementation, you can save the current status of your work in your repository following the submission instructions below. When you continue working on the assignment, do this in the version of the notebook that you saved in your Google Drive (access from your Google Colab account) or your repository.
4. After finishing the assignment submit your assignment as follows:
    * Check that everything runs as expected. To do that select in the menubar ``Runtime`` $\rightarrow$ ``Restart and run all``. This will clear all your local variables and runs your notebook from the beginning
    * Save your notebook, click ``File`` $\rightarrow$ ``Save``
    * Click ``File`` $\rightarrow$ ``Save a copy in GitHub``
    * Select ``digikolleg-data-science-psychology/05_dicts_classes-<username>`` from the Repository drop-down menu 
    * Type ``05_dicts_classes.ipynb`` into the File path text box
    * Optionally, you may enter a description into the "Commit message" text box (like: "solved task 1" or "ready for grading")
    * Make sure that the "Include a link to Colaboratory" box is checked
    * Press "Ok"
    * A new tab or window should open up, displaying your submitted notebook within your GitHub repository

Before the assignment deadline, you can re-submit your assignment as often as you would like. This will update the copy of your assignment file in your assignment's repository.

### **Assignment Description**
In this assignment, you will explore two essential concepts in Python. *Dictionaries* are data structures that allow you to store, retrieve, and group data. *Classes* are the building blocks of object-oriented programming (OOP) and provide a way to create reusable, modular code. Learning *dictionaries* and *classes* fosters an object-oriented mindset, which is a powerful and widely-used programming paradigm.

#### **Assignment goals**
You are going to learn the following things in this assignment:
1. how to create a dictionary and add, remove, and change its elements
2. how to create classes and instances of classes
3. how to write class-specific methods
4. how to use class inheritance
5. how to define relationships between objects

#### **Assignment point value**
This assignment is worth a total (up to) **13** points. If your code passes all of the manual and automatic tests you will receive full credit.

### **Tasks**
Make sure that you now work on your *copied* version of this assignment's notebook.

#### 0. Imports and Setup
Please **run** the cell below to import the libraries needed for this assignment.

```python
from datetime import date, datetime
from contextlib import suppress
```

#### 1. Dictionary
A dictionary is a type of data structure that stores key-value pairs. Each key in the dictionary is associated with a value:
```
dict_1 = {
    key_1: value_1,
    key_2: value_2
}
```

##### 1.1 Create a dictionary called ``patient`` with the following key-value pairs:

| key | value |
| :- | :- |
|name|Jane|
|height|167.3|
|vegetarian|True|
|allergies|[pollen, peanuts]|
|blood_test_results|None|

<!-- grade_id: create_dict, points: 0, grade: false, solution: true -->
```python
patient= {
  'name': 'Jane',
    'height': '167.3',
    'vegetarian': 'True',
    'allergies': '[pollen, peanuts]',
    'blood_test_results': 'None'
}
```

<!-- grade_id: create_dict_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert patient["name"] == "Jane", "Check if you assigned the correct value to the keys"
assert len(patient) == 5, "Check if you added all key-value pairs"
```

##### 1.2 Add the key ``age`` with value 24 to the ``patient`` dictionary:

<!-- grade_id: add_key, points: 0, grade: false, solution: true -->
```python
patient['age']= 24
print(patient)
```

Output:
```
{'name': 'Jane', 'height': '167.3', 'vegetarian': 'True', 'allergies': '[pollen, peanuts]', 'blood_test_results': 'None', 'age': 24}
```

<!-- grade_id: add_key_test, points: 0.5, grade: true, solution: false -->
```python
# run test to check your implementation
assert patient["age"], "Check if you added the key-value pair correctly"
```

##### 1.3 Create another dictionary called ``blood_test_results`` and change the ``blood_test_results`` key of ``patient``:
| key | value |
| :- | :- |
|hemoglobin|13.6|
|potassium|4.5|

<!-- grade_id: change_dict, points: 0, grade: false, solution: true -->
```python
patient['blood_test_results']= {
  'hemoglobin': '13.6',
    'potassium': '4.5'
}
print(patient)
```

Output:
```
{'name': 'Jane', 'height': '167.3', 'vegetarian': 'True', 'allergies': '[pollen, peanuts]', 'blood_test_results': {'hemoglobin': '13.6', 'potassium': '4.5'}, 'age': 24}
```

<!-- grade_id: change_dict_test, points: 0.5, grade: true, solution: false -->
```python
# run test to check your implementation
assert len(patient["blood_test_results"]) == 2, "Check if you added both elements to the dictionary"
```

##### 1.4 Remove the key-value pair associated with ``allergies`` from ``patient``:

<!-- grade_id: remove_key, points: 0, grade: false, solution: true -->
```python
patient.pop('allergies')
print(patient)
```

Output:
```
{'name': 'Jane', 'height': '167.3', 'vegetarian': 'True', 'blood_test_results': {'hemoglobin': '13.6', 'potassium': '4.5'}, 'age': 24}
```

<!-- grade_id: remove_key_test, points: 0.5, grade: true, solution: false -->
```python
# run test to check your implementation
with suppress(KeyError): patient["allergies"]; assert False, "Ensure that you removed the correct key-value pair"
```

#### 2. Classes - Medical records system
In this task, you will create a simulation of a patient's medical records system. Therefore, you will create 5 different classes:
* ``Medication``: Stores information like name and dosage for one type of medication. It is possible to change the prescribed dosage.
* ``Session``: Contains the information about any appointment a patient could have. Also, notes can be stored after the session.
* ``Physician``: Is the base class for every physician. A physician can schedule sessions for their patients and add notes to the sessions.
* ``PrimaryCarePhysician``: Inherits all functionalities from ``Physician`` and is able to refer a patient to a specialist.
* ``Patient``: Is the base of the medical records system. All medications, sessions, and referrals are stored there. Also, any modifications are initiated using this class. Furthermore, it can visualize all stored entries for the user.

##### 2.1 Create the class ``Medication``, to do so:
* Create a class ``Medication``
* Create the constructor (``__init__()``) with ``name`` and ``dosage`` as parameters. Initialize the attributes of the object ``self.name`` and ``self.dosage`` with the respective parameter.
* Create the method ``update_dosage`` with ``new_dosage`` as parameter which updates the object attribute ``self.dosage``
* Provide a meaningful string representation of your object using the ``__str__(self)`` method (one of the [special methods](https://rszalski.github.io/magicmethods/#representations) in Python). For example, if the object's ``name`` is "Ibuprofen" and ``dosage`` is "200 mg" you could return: "Medication: Ibuprofen, Dosage: 200 mg". You can use [Formatted String Literals](https://docs.python.org/3/tutorial/inputoutput.html#tut-f-strings) for that.
* Create an instance of ``Medication`` to test your implementation and try to access its object attributes

**Hint**: If the instantiation doesn't work properly, have a look at the ``self`` keyword and its meaning and usage.

<!-- grade_id: medication, points: 0, grade: false, solution: true -->
```python
class Medication:

  def  __init__(self,name,dosage):
    self.name=name
    self.dosage=dosage

  def update_dosage(self,new_dosage):
    self.dosage= new_dosage

  def __str__(self):
    return f"Medication: {self.name}, Dosage: {self.dosage}"

medication = Medication
```

<!-- grade_id: medication_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
medication = Medication("Ibuprofen", "200 mg")
assert medication.name == "Ibuprofen" and  medication.dosage == "200 mg", "Ensure that you save the constructor parameters as object attributes"
medication.update_dosage("400 mg"),
assert medication.dosage == "400 mg", "Ensure that you save new_dosage in the respective object attribute"
assert medication.name in str(medication) and medication.dosage in str(medication), "Ensure that you return a string with all the relevant information"
```

##### 2.2 Create the class ``Session``, to do so:
* Create a class ``Session``
* Create a constructor with ``session_date``, ``physician_name``, and ``purpose`` as parameters
    * Save ``session_date`` only as object attribute if it is of type ``datetime.date``. Otherwise, raise a TypeError with a meaningful message for the user. The class ``date`` is already imported from the module ``datetime`` at the beginning of the assignment.
    * Save ``physician_name`` only as object attribute if it is of type ``str``. Otherwise raise a TypeError with a meaningful message.
    * Save ``purpose`` as an object attribute with the name accordingly
    * Furthermore, create an object attribute ``notes`` and initialize it with an empty dictionary.
* Create a method ``add_session_notes`` with the parameter ``notes``. Save ``notes`` in the respective object attribute using the current local date and time as key. Check out the [datetime](https://docs.python.org/3/library/datetime.html#module-datetime) module for a suitable function.
* Provide a string representation of ``Session`` like for ``Medication``. Include all object attributes (also the notes). **Hint**: ``date`` also has a built-in string representation.
* Create an instance of ``Session`` and test your methods

<!-- grade_id: session, points: 0, grade: false, solution: true -->
```python
class Session:
    def __init__(self, session_date, physician_name, purpose):
        if not isinstance(session_date, date):
            raise TypeError("session_date not in right format")
        self.session_date = session_date

        if not isinstance(physician_name, str):
            raise TypeError("physician_name must be a str")
        self.physician_name = physician_name

        self.purpose = purpose
        self.notes = {}

    def add_session_notes(self, notes):
        present_datetime = datetime.now()
        self.notes[present_datetime] = notes

    def __str__(self):
        session_str = f"Session Date: {self.session_date}\n Physician Name: {self.physician_name}\n Purpose: {self.purpose}\n Notes:\n"
        for session_info, notes in self.notes.items():
            session_str += f"{session_info}: {notes}\n"
            print(session_str)
        return session_str
```

<!-- grade_id: session_test, points: 2, grade: true, solution: false -->
```python
# run tests to check your implementation
session = Session(date(year=2000, month=1, day=1), "Dr. X", "blood_test")
assert session.session_date == date(year=2000, month=1, day=1) and session.physician_name == "Dr. X" and session.purpose == "blood_test" and len(session.notes) == 0, "Ensure that you set all object attributes correctly"
with suppress(TypeError): Session("2000, 1, 1", None, None); assert False, "Ensure that you raise the correct type of Exception"
with suppress(TypeError): Session(date(year=2000, month=1, day=1), None, None); assert False, "Ensure that you raise the correct type of Exception"
session.add_session_notes("anxiety")
assert set(map(type, session.notes)) == {datetime} and set(session.notes.values()) == {"anxiety"}, "Ensure that you used datetime as key and the given notes as values"
assert str(session.session_date) in str(session) and session.physician_name in str(session) and session.purpose in str(session) and "anxiety" in str(session), "Ensure that you use all relevant information in your string representation"
```

Output:
```
Session Date: 2000-01-01
 Physician Name: Dr. X
 Purpose: blood_test
 Notes:
2023-06-06 08:05:49.736785: anxiety

Session Date: 2000-01-01
 Physician Name: Dr. X
 Purpose: blood_test
 Notes:
2023-06-06 08:05:49.736785: anxiety

Session Date: 2000-01-01
 Physician Name: Dr. X
 Purpose: blood_test
 Notes:
2023-06-06 08:05:49.736785: anxiety

Session Date: 2000-01-01
 Physician Name: Dr. X
 Purpose: blood_test
 Notes:
2023-06-06 08:05:49.736785: anxiety
```

##### 2.3 Create the class ``Patient`` with all its three subparts **medication handling**, **session handling**, and **referral handling**, to do so:
* Create a class ``Patient``
* Create the constructor which takes ``name`` as parameter and saves it as an object attribute
    * Initialize object attributes ``medications``, ``sessions``, and ``referrals`` as empty dictionaries
    * Initialize an object attribute ``session_id`` with 0

**Medication handling:**
* Create a method ``add_medication`` which takes an object of ``Medication`` as parameter ``medication``
    * Add ``medication`` to the respective object attribute using its name as key
* Create a method ``remove_medication`` with ``medication_name`` as parameter
    * Remove the medication from the respective object attribute
* Create a method ``update_medication_dosage`` with ``medication_name`` and ``new_dosage`` as parameters
    * Update the dosage of the given medication
* Create a method ``print_medications`` which prints all the prescribed medications for this patient using the defined string representation of ``Medication``

**Session handling:**
* Create a methode ``increment_session_id`` which increments ``session_id`` by 1
* Create a method ``add_therapy_session`` with ``session`` as parameter
    * Perform a type check of ``session``. If it is of type ``Session`` add ``session`` to the respective object attribute using ``session_id``. Otherwise, raise a TypeError with a meaningful message. 
    * Note: ``session_id`` must be unique, first session gets an id equal to 0
* Create a method ``get_session`` with ``session_id`` as parameter which returns the corresponding ``Session`` object
* Create a method ``print_sessions`` which prints all sessions for this patient using the defined string representation of ``Session``

**Referral handling:**
* Create a method ``refer_to_specialist`` with ``specialist`` (of type str) as parameter
    * Add ``specialist`` to the respective object attribute using the current local date and time as key
* Create a method ``print_referrals`` which prints all referrals for this patient

**Testing:**
* Create an instance of patient
* Use the different methods to test your implementation
* Try your ``print_...`` methods. Do they show all relevant information? If not check your string representations of ``Medication`` and ``Session``.

<!-- grade_id: patient, points: 0, grade: false, solution: true -->
```python

from datetime import datetime

class Patient:
    def __init__(self, name):
        self.name = name
        self.medications = {}
        self.sessions = {}
        self.referrals = {}
        self.session_id = 0

    def add_medication(self, medication):
        self.medications[medication.name] = medication

    def remove_medication(self, medication_name):
        if medication_name in self.medications:
            del self.medications[medication_name]

    def update_medication_dosage(self, medication_name, new_dosage):
        if medication_name in self.medications:
            self.medications[medication_name].update_dosage(new_dosage)

    def print_medications(self):
        for medication in self.medications.values():
            print(medication)

    def increment_session_id(self):
        self.session_id += 1

    def add_therapy_session(self, session):
        if isinstance(session, Session):
            self.sessions[self.session_id] = session
            self.increment_session_id()
        else:
            raise TypeError("session must be of type Session")

    def get_session(self, session_id):
        return self.sessions.get(session_id)

    def print_sessions(self):
        for session_id, session in self.sessions.items():
            print(f"Session ID: {session_id}")
            print(session)

    def refer_to_specialist(self, specialist):
        current_datetime = datetime.now()
        self.referrals[current_datetime] = specialist

    def print_referrals(self):
        for referral_datetime, specialist in self.referrals.items():
            print(f"{referral_datetime}: {specialist}")


class Medication:
    def __init__(self, name, dosage):
        self.name = name
        self.dosage = dosage

    def update_dosage(self, new_dosage):
        self.dosage = new_dosage

    def __str__(self):
        return f"Medication: {self.name}, Dosage: {self.dosage}"


class Session:
    def __init__(self, session_date, physician_name, purpose):
        self.session_date = session_date
        self.physician_name = physician_name
        self.purpose = purpose
        self.notes = {}

    def add_session_notes(self, notes):
        current_datetime = datetime.now()
        self.notes[current_datetime] = notes

    def __str__(self):
        session_str = f"Session Date: {self.session_date}\nPhysician Name: {self.physician_name}\nPurpose: {self.purpose}\nNotes:\n"
        for session_datetime, notes in self.notes.items():
            session_str += f"{session_datetime}: {notes}\n"
        return session_str
```

Run tests to check your implementation for **medication handling**:

<!-- grade_id: patient_medication_test, points: 2, grade: true, solution: false -->
```python
patient = Patient("Jane Doe")
assert patient.name == "Jane Doe" and len(patient.medications) == 0 and len(patient.sessions) == 0 and len(patient.referrals) == 0 and patient.session_id == 0, "Ensure that you set all object attributes correctly"
med = Medication("Ibuprofen", "200 mg")
patient.add_medication(med)
assert len(patient.medications) == 1 and set(patient.medications.keys()) == {med.name} and patient.medications[med.name] == med, "Ensure that you used the medication name as key and the Medication object as value"
patient.update_medication_dosage(med.name, "400 mg")
assert len(patient.medications) == 1 and patient.medications[med.name].dosage == "400 mg", "Ensure that you used the Medication method update_dosage"
patient.remove_medication(med.name)
assert len(patient.medications) == 0, "Check if you delete the Medication object in remove_medication"
```

Run tests to check your implementation for **session handling**:

<!-- grade_id: patient_session_test, points: 2, grade: true, solution: false -->
```python
patient = Patient("Jane Doe")
patient.increment_session_id()
assert patient.session_id == 1, "Check if your stored the incremented session_id as object attribute"
session = Session(date(year=2000, month=1, day=1), "Dr. X", "blood_test")
patient.add_therapy_session(session)
assert len(patient.sessions) == 1 and set(patient.sessions.keys()) == {1} and patient.sessions[1] == session, "Ensure that you added the session correctly in add_therapy_session"
assert patient.session_id == 2, "Check if you incremented session_id after session adding"
assert patient.get_session(1) == session, "Ensure that you return the correct Session object"
```

Run tests to check your implementation for **referral handling**:

<!-- grade_id: referral_test, points: 1, grade: true, solution: false -->
```python
patient = Patient("Jane Doe")
patient.refer_to_specialist("psychologist")
assert len(patient.referrals) == 1 and set(map(type, patient.referrals.keys())) == {datetime} and set(patient.referrals.values()) == {"psychologist"}, "Ensure that you added the referral correctly"
```

##### 2.4 Create the class ``Physician``, to do so:
* Create a class ``Physician``
* Create the constructor which takes ``name`` as parameter and stores it as an object attribute
* Create the method ``schedule_session`` with ``session_date``, ``patient``, and ``purpose`` as parameters
    * Create an instance of ``Session`` with the given information
    * Add the session to the therapy sessions of ``patient``
* Create the method ``add_session_notes`` with ``patient``, ``session_id``, and ``notes`` as parameters
    * Use the ``get_session`` method of ``patient`` to get the respective ``Session`` object of ``session_id``
    * add ``notes`` to the session

<!-- grade_id: physician, points: 0, grade: false, solution: true -->
```python

class Physician:
    def __init__(self, name):
        self.name = name

    def schedule_session(self, session_date, patient, purpose):
        session = Session(session_date, self.name, purpose)
        patient.add_therapy_session(session)

    def add_session_notes(self, patient, session_id, notes):
        session = patient.get_session(session_id)
        if session:
            session.add_session_notes(notes)


class Session:
    def __init__(self, session_date, physician_name, purpose):
        self.session_date = session_date
        self.physician_name = physician_name
        self.purpose = purpose
        self.notes = {}

    def add_session_notes(self, notes):
        current_datetime = datetime.now()
        self.notes[current_datetime] = notes

    def __str__(self):
        session_str = f"Session Date: {self.session_date}\nPhysician Name: {self.physician_name}\nPurpose: {self.purpose}\nNotes:\n"
        for session_datetime, notes in self.notes.items():
            session_str += f"{session_datetime}: {notes}\n"
        return session_str
```

<!-- grade_id: physician_test, points: 2, grade: true, solution: false -->
```python
# run tests to check your implementation
physician = Physician("Dr. X")
assert physician.name == "Dr. X"
patient = Patient("Jane Doe")
physician.schedule_session(date(2000, 1, 1), patient, "blood_test")
assert len(patient.sessions) == 1, "Ensure that you use the patient's add_therapy_session method"
physician.add_session_notes(patient, 0, "patient reported anxiety")
session = patient.get_session(0)
assert set(session.notes.values()) == {"patient reported anxiety"}, "Ensure that you used the add_session_notes of session"
```

##### 2.5 Create the class ``PrimaryCarePhysician``, to do so:
* Create a class ``PrimaryCarePhysician`` that inherits from ``Physician``
* Create a method ``refer_to_specialist`` with ``patient`` and ``specialist`` (type str) as parameters
    * Use the ``refer_to_specialist`` method from ``Patient`` to add the referral to ``patient``

<!-- grade_id: prim_care_physician, points: 0, grade: false, solution: true -->
```python
class PrimaryCarePhysician(Physician):
    def refer_to_specialist(self, patient, specialist):
        patient.refer_to_specialist(specialist)
```

<!-- grade_id: prim_care_physician_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
prim_care_phys = PrimaryCarePhysician("Dr. Z")
patient = Patient("Jane Doe")
prim_care_phys.schedule_session(date(2000, 1, 1), patient, "blood_test")
assert len(patient.sessions) == 1, "Ensure that PrimaryCarePhysician inherits all properties from Physician"
prim_care_phys.refer_to_specialist(patient, "psychologist")
assert len(patient.referrals) == 1, "Ensure that you used the patient's refer_to_specialist method"
```

##### 2.6 Play around with the patient's medical records system that you created:
* Create one or more patients
* Create different physicians and schedule sessions for your patients
* Add the session notes
* Add, remove, and change your patients' medications
* refer your patients to different specialists
* Use your patients' visualization system (``print_medications()``... methods) to keep track of all the changes

<!-- grade_id: playground, points: 0, grade: false, solution: true -->
```python

#patients
patient1 = Patient("Adam")
patient2 = Patient("Bob")

#physicians
physician1 = PrimaryCarePhysician("Dr. Oswald")
physician2 = Physician("Dr. Patel")

#sessions for patients
physician1.schedule_session(date(2023, 5, 5), patient1, "general consultation")
physician2.schedule_session(date(2023, 6, 6), patient2, "reports checking")

#session notes
physician1.add_session_notes(patient1, 0, "Patient shows no symptoms of Diabetis.")
physician2.add_session_notes(patient2, 0, "Patient reports are normal.")

#medications for patients
medication1 = Medication("Ibuprofen", "200 mg")
medication2 = Medication("Ibuflam", "400 mg")
patient1.add_medication(medication1)
patient2.add_medication(medication2)

# Print medications for patients
patient1.print_medications()
patient2.print_medications()

# Update medication dosage
patient1.update_medication_dosage("Ibuprofen", "400 mg")

# Remove medication
patient2.remove_medication("Ibuflam")

# Refer patients to specialists
patient.refer_to_specialist("Endocrinologist")
patient.refer_to_specialist("Orthopaedic")

# Print sessions for patients
patient1.print_sessions()
patient2.print_sessions()

# Print referrals for patients
patient1.print_referrals()
patient2.print_referrals()
```

Output:
```
Medication: Ibuprofen, Dosage: 200 mg
Medication: Ibuflam, Dosage: 400 mg
Session ID: 0
Session Date: 2023-05-05
Physician Name: Dr. Oswald
Purpose: general consultation
Notes:
2023-06-06 08:05:49.851895: Patient shows no symptoms of Diabetis.

Session ID: 0
Session Date: 2023-06-06
Physician Name: Dr. Patel
Purpose: reports checking
Notes:
2023-06-06 08:05:49.851939: Patient reports are normal.
```

```python

```