<a href="https://colab.research.google.com/github/suchirasuresh/Applied-Data-Science-in-Medicine-Psychology-ads_assignment_1--suchira_suresh-/blob/main/ads_assignment_3.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Fill in your github username and idm credentials below:

```python
github_username = "suchirasuresh"
idm_credentials = "yh07ajaw"
```

---

# **Applied Data Science in Medicine & Psychology - Assignment 3**
---
**Deadline: 16.05.2023, 23:59**
### **General setup and submission**
1. Save an editable copy of this notebook to your Google Drive ``File`` $\rightarrow$ ``Save a copy in Drive`` and remove the automatically appended "Copy of" text. **Do not change the file name**.
2. Work on the assignment in the *copied* version of the notebook. Make sure that you **replace** all parts of the tasks indicated with *YOUR CODE HERE* and **raise NotImplementedError()** (otherwise this causes an error in your implementation) or *YOUR ANSWER HERE* with your solution.
3. If you do a break during the implementation, you can save the current status of your work in your repository following the submission instructions below. When you continue working on the assignment, do this in the version of the notebook that you saved in your Google Drive (access from your Google Colab account) or your repository.
4. After finishing the assignment submit your assignment as follows:
    * Check that everything runs as expected. To do that select in the menubar ``Runtime`` $\rightarrow$ ``Restart and run all``. This will clear all your local variables and runs your notebook from the beginning
    * Save your notebook, click ``File`` $\rightarrow$ ``Save``
    * Click ``File`` $\rightarrow$ ``Save a copy in GitHub``
    * Select ``digikolleg-data-science-psychology/ads_assignment_3-<username>`` from the Repository drop-down menu 
    * Type ``ads_assignment_3.ipynb`` into the File path text box
    * Optionally, you may enter a description into the "Commit message" text box (like: "solved task 1" or "ready for grading")
    * Make sure that the "Include a link to Colaboratory" box is checked
    * Press "Ok"
    * A new tab or window should open up, displaying your submitted notebook within your GitHub repository

Before the assignment deadline, you can re-submit your assignment as often as you would like. This will update the copy of your assignment file in your assignment's repository.

### **Assignment Description**
In this assignment, we cover the second part of Python basics. You will learn essential programming building blocks such as functions, data types and structures, control flow elements, and data type casting which are used to store, manipulate, and control data in a program. Understanding these concepts is crucial for writing efficient and effective code.

#### **Assignment goals**
You are going to learn the following things in this assignment:
1. how to implement and use functions
2. how to create variables of different data types
3. how to set up and work with dictionaries and lists
4. how to control the program flow with if-statements and loops
5. how to cast variables to different datatypes

#### **Assignment point value**
This assignment is worth a total (up to) **10** points. If your code passes all of the manual and automatic tests you will receive full credit.

### **Tasks**
Make sure that you now work on your *copied* version of this assignment's notebook.

#### 1. Functions
Functions are a set of instructions that perform a specific task within a program. They allow developers to modularize code and reuse it throughout the program. Functions can take arguments (inputs) and return values (outputs):
```
def <function_name>(<arg_1>, ..., <arg_N>):
    <instruction_1>
    ...
    <instruction_N>
    return <value>
```

##### 1.1 Create a function called ``get_datatype`` that takes a variable called ``var`` as argument and returns the type of the variable. Test your solution with variables of different datatypes:
**Hint:** check the lecture slides for a suitable function.

<!-- grade_id: function_datatype, points: 0, grade: false, solution: true -->
```python
def get_datatype(var):
  return type(var)
```

<!-- grade_id: function_datatype_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert get_datatype(5) == int, "Check if you return the data type itself and not a string representation or similar"
assert get_datatype("Hello, world!") == str, "Check if you return the data type itself and not a string representation or similar"
```

#### 2. Datatypes - List
A Python list is a structure that allows you to store a collection of elements of any data type in a single variable, like
```
python_list = [12, "hello", 23.5, 25]
```
Lists are provided with the following functionalities:

| functionality | key word|explanation | example | return |
| :- | :- | :- | :- | :- |
|Accessing||access individual elements by their index, <br /> index starts from 0|``python_list[0]``|``12``|
|Slicing||create a sub-list, <br /> syntax: ``lst[start:stop:step]``|``python_list[1:4:2]`` <br /> |``["hello", 25]``|
|Inserting|``insert``|insert a new element at a specific position|``python_list.insert(2, "new")`` | ``[12, "hello", "new", 23.5, 25]``
|Deleting|``remove``|remove an element|``python_list.remove("hello")``|``[12, 23.5, 25]``|
|Appending|``append``|add new elements to the end|``python_list.append([26, 27])``|``[12, "hello", 23.5, 25, [26, 27]]``|
|Extending|``extend``|add new elements to the end|``python_list.extend([26, 27])``|``[12, "hello", 23.5, 25, 26, 27]``|
|Membership|``in``|check if a particular element is present|``25 in python_list`` <br /> ``56 in python_list``|``True`` <br /> ``False``

##### 2.1 Create a list called ``my_list`` with the elements ``1, 2, "hello", 4, True``:

<!-- grade_id: list, points: 0, grade: false, solution: true -->
```python
my_list = [1, 2, "hello", 4, True]
```

<!-- grade_id: list_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert my_list
assert len(my_list) == 5, "Check if you created the list with all the elements"
assert my_list[2] == "hello", "Check if all elements are in the right order"
```

##### 2.2 Create a variable ``result`` which adds the 2nd and 4th element of ``my_list`` using their index:

<!-- grade_id: list_index, points: 0, grade: false, solution: true -->
```python
result = my_list[1]+my_list[3]
```

<!-- grade_id: list_index_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert result
assert result == 6
```

##### 2.3 Create a list ``sub_list`` which contains the 3rd, 4th, and 5th element of ``my_list`` using slicing:

<!-- grade_id: list_slicing, points: 0, grade: false, solution: true -->
```python
sub_list= my_list[2:5:1]
print (sub_list)
```

Output:
```
['hello', 4, True]
```

<!-- grade_id: list_slicing_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert sub_list
assert len(sub_list) == 3, "Check the length of your sub_list and your indices. Did you consider how the last element is handled in slicing?"
```

##### 2.4 Insert at the 4th position of ``new_list`` the string ``world``. Delete ``4`` out of ``new_list``:

<!-- grade_id: list_ins_del, points: 0, grade: false, solution: true -->
```python
new_list = my_list.copy()
new_list.insert(4, "world")
new_list.remove(4)


print(new_list)
```

Output:
```
[1, 2, 'hello', 'world', True]
```

<!-- grade_id: list_ins_del_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert len(new_list) == 5, "Check if you performed both operations (inserting, deleting)"
assert new_list[3] == "world", "Check if you used the correct index for inserting"
assert 4 not in new_list, "Check if you deleted 4 out of the list"
```

##### 2.5 Append ``tmp_list`` to ``new_list``:

<!-- grade_id: list_append, points: 0, grade: false, solution: true -->
```python
tmp_list = [False, "hi", 3]
new_list = my_list.copy()
new_list.append(tmp_list)



print(new_list)
print(f"length of new_list: {len(new_list)}")
```

Output:
```
[1, 2, 'hello', 4, True, [False, 'hi', 3]]
length of new_list: 6
```

<!-- grade_id: list_append_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert len(new_list) == 6, "The length of new_list must be 6"
```

##### 2.6 Extend ``new_list`` with ``tmp_list`` from above:

<!-- grade_id: list_extend, points: 0, grade: false, solution: true -->
```python
new_list = my_list.copy()

new_list.extend(tmp_list)

print(new_list)
print(f"length of new_list: {len(new_list)}")
```

Output:
```
[1, 2, 'hello', 4, True, False, 'hi', 3]
length of new_list: 8
```

<!-- grade_id: list_extend_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert len(new_list) == 8, "The length of new_list must be 8"
```

* Can you explain the difference between ``append`` and ``extend``?



```
Append: Adds one element at the end of the parent list. Argument passed in the append function is added as single element at end of the list and the length of the list is increased by 1 only.

Extend: Adds each element of the iterable passed to the end of the parent list . It increases the length of the list by the number of elements of the iterable passed as an argument.
```



##### 2.7 Create 2 variables
* ``var_1``: checks whether ``my_list`` contains "hello, world!"
* ``var_2``: checks whether ``my_list`` contains 4

<!-- grade_id: list_member, points: 0, grade: false, solution: true -->
```python
var_1 = "hello, world!" in my_list
var_2 = 4 in my_list


print(var_1)
print(var_2)
```

Output:
```
False
True
```

<!-- grade_id: list_member_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert type(var_1) == bool
assert type(var_2) == bool
```

#### 3. if-statement
An if-statement is a control flow statement that allows you to execute a certain block of code only if a particular condition evaluates to true.
##### 3.1 Create an if - elif - else statement that checks whether a variable is equal to, greater, or less than 0. To do so:
* create a variable called ``x`` and assign -5 to it
* if x is equal to 0: print "x is equal to 0"
* if x is less than 0: print "x is negative"
* if x is greater than 0: print "x is positive"
* play around with different values for x to check your implementation

<!-- grade_id: simple_if, points: 0, grade: false, solution: true -->
```python
x = -5
if x==0:
  print ("x is equal to 0")
elif x<0:
  print("x is negative")
else:
  print("x is positive")



```

Output:
```
x is negative
```

##### 3.2 Implement a function that checks the potassium level of a person. The potassium level is usually between 3.7 and 5.3 mEq/l (milliequivalents per liter), inclusive both endpoints. To do so:
* create a function called ``check_potassium`` that takes ``value`` as an argument
* check if ``value`` is in the given normal potassium range
    * if yes: print an informative message (e.g. "the potassium level is in the correct range") and return True
    * if no: print an informative message and return False
* Play around with different values to see if your implementation works correctly

**Hint:** You can combine two conditions using logical operators. Also, you can create a variable to store your return value temporarily.

<!-- grade_id: check_potassium, points: 0, grade: false, solution: true -->
```python
def check_potassium (value):
  #value = 4.5
  if (value >=3.7 and value <= 5.3):
    print("potassium level is in the correct range")
    out = True
  else:
    print("potassium level is abnormal")
    out = False
  return out

```

<!-- grade_id: check_potassium_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert check_potassium(4.0) == True, "Check if you used the correct potassium range"
assert check_potassium(3.0) == False, "Check if you used the correct potassium range"
assert check_potassium(5.3) == True, "Check if you used the correct potassium range"
```

Output:
```
potassium level is in the correct range
potassium level is abnormal
potassium level is in the correct range
```

##### 3.3 Implement a function that checks the hemoglobin level based on the age of the person. Assume the hemoglobin levels to be as follows:
* 0 - 17 years: 10 - 15.5 g/dl (grams per deciliter)
* \>= 18 years: 12 - 18 g/dl (grams per deciliter)

To do so:
* create a function called ``check_hemoglobin`` that takes ``age`` and ``value`` as arguments
* create a nested if-construct
    * outer if-statement checks the range of ``age``
    * inner if-statements checks based on the age the correct range of hemoglobin level
        * if it is in the correct range print an informative message and return True
        * if it is **not** in the correct range print an informative message and return False



<!-- grade_id: check_hemoglobin, points: 0, grade: false, solution: true -->
```python
def check_hemoglobin(age, value):
  if age >= 18:
    if value>=12.0 and value<=18.0:
      print("reports of the adult are normal")

    else:
      print("reports of the adult are abnormal")

  else:
    if value>=10.0 and value<=15.5:
      print("reports of the child are normal")

    else:
      print("reports of the child are abnormal")


```

<!-- grade_id: check_hemoglobin_test, points: 2, grade: true, solution: false -->
```python
# run tests to check your implementation
assert check_hemoglobin(18, 10.5) == False, "Check if you set the age range correctly. Check if you used the correct hemoglobin range for adult people"
#assert check_hemoglobin(23, 12.0) == True, "Check if you used the correct hemoglobin range for adult people"
#assert check_hemoglobin(23, 11.9) == False, "Check if you used the correct hemoglobin range for adult people"
#assert check_hemoglobin(16, 15.5) == True, "Check if you used the correct hemoglobin range for underage people"
#assert check_hemoglobin(16, 15.6) == False, "Check if you used the correct hemoglobin range for underage people"
```

Output:
```
reports of the adult are abnormal
AssertionError                            Traceback (most recent call last)
<ipython-input-91-bff26848ff32> in <cell line: 2>()
      1 # run tests to check your implementation
----> 2 assert check_hemoglobin(18, 10.5) == False, "Check if you set the age range correctly. Check if you used the correct hemoglobin range for adult people"
      3 #assert check_hemoglobin(23, 12.0) == True, "Check if you used the correct hemoglobin range for adult people"
      4 #assert check_hemoglobin(23, 11.9) == False, "Check if you used the correct hemoglobin range for adult people"
      5 #assert check_hemoglobin(16, 15.5) == True, "Check if you used the correct hemoglobin range for underage people"

AssertionError: Check if you set the age range correctly. Check if you used the correct hemoglobin range for adult people
```

#### 4. Loops
Loops are a way to execute a block of code repeatedly. There are typically two types of loops:
1. **For loops** allow you to iterate over a sequence of elements (such as a list or range of numbers) and execute a block of code for each element in the sequence.
2. **While loops** allow you to execute a block of code repeatedly while a certain condition is true.

##### 4.1 Create a function ``calc_sum`` that takes a list of values as argument ``values``. Then calculate the sum of the list entries using a for loop and return the sum:
**Hint:** Initialize a variable ``sum`` before the loop and assign it to 0. Test your solution with ``numbers``.

<!-- grade_id: for_sum, points: 0, grade: false, solution: true -->
```python
numbers = [23, 27, 21, 19]

def calc_sum(values):
  sum=0
  #values = numbers
  for i in values:
    sum+=i
  print(sum)
  return sum

```

<!-- grade_id: for_sum_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation 
assert calc_sum([23, 27, 21, 19]) == 90, "Check your for loop syntax and used operator"
```

Output:
```
90
```

##### 4.2 Create a function ``arange`` which generates lists of ascending numbers using a while loop. It takes the following arguments:
* ``start``: start of the interval, included
* ``stop``: stop of the interval, excluded
* ``step``: spacing between values

To do so:
* create a control variable ``n`` and initialize it with ``start``
* initialize an empty list ``numbers``
* create a while loop that checks the current value of ``n`` and runs until ``stop``
* append ``n`` to ``numbers``
* add ``step`` in every iteration to your control variable.
* return ``numbers``

<!-- grade_id: while_arange, points: 0, grade: false, solution: true -->
```python
def arange(start,stop,step):
  n= start
  numbers=[]
  while (n<stop):
    numbers.append(n)
    n+=step
  print (numbers)

  return numbers


```

<!-- grade_id: while_arange_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert arange(5, 10, 1) == [5, 6, 7, 8, 9], "Check if you set start and stop correctly, used the right operator, and performed all steps in the right order"
assert arange(5, 10, 2) == [5, 7, 9], "Check if you set start and stop correctly, used the right operator, and performed all steps in the right order"
```

Output:
```
[5, 6, 7, 8, 9]
[5, 7, 9]
```

#### 5. Typecasting
Typecasting, also known as type conversion, is the process of converting a value of one data type to another data type. This can be useful in programming when you want to perform operations that require values of a specific data type.

##### 5.1 Convert the following variables to the specified data types:
* x from ``float`` to ``int``
* y from ``float`` to ``bool``
* z from ``int`` to ``float``

<!-- grade_id: typecasting, points: 0, grade: false, solution: true -->
```python
x = 8.7
y = 1.0
z = 9

x = int(x)
y = bool(y)
z = float(z)

print(x)
print(y)
print(z)
```

Output:
```
8
True
9.0
```

<!-- grade_id: typecasting_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert type(x) == int, "Check if you converted x to the correct data type"
assert type(y) == bool, "Check if you converted y to the correct data type"
assert type(z) == float, "Check if you converted z to the correct data type"
```