# **Applied Data Science in Medicine & Psychology** - Assignment 1
---
**Deadline: 25.04.2023, 23:59**

This first assignment of the "Applied Data Science in Medicine & Psychology" course introduces the tools used for the programming tasks and gives instructions about their setup.

### **Setup**  
1. Click on the blue "Open in Colab" button on the top of the assignment. A new tab or window should open in your browser, displaying this notebook in Google Colaboratory.
2. **Continue the setup of your assignment whithin Google Colaboratory with the following point**
3. To be able to edit and submit the assignment you need to save an editable copy of your notebook. Click ``File`` $\rightarrow$ ``Save a copy in Drive`` and then save your copy as ``ads_assignment_1`` (i.e. remove the "Copy of" text that gets automatically appended to the beginning of the filename). Do not change the file name.


### **Assignment Description**

#### **Assignment goals**
You are going to learn the following things in this assignment:
1. How to setup your Google Colaboratory and GitHub accounts
2. How to access and submit assignments

#### **Assignment point value**
This assignment is worth a total (up to) **1** point. If your code passes all of the automatic tests, you will receive full credit.

### **Tasks**
Make sure that you now work on your *copied* version of this assignment's notebook.

1. Change the code in the next cell to indicate that you have gone through the exercise introduction in StudOn and set up your Google Colaboratory and GitHub account. You should change each variable from ``False`` to ``True`` to show which tasks you have accomplished (by self-report). By running the cell you can check for errors in your implementation. Make sure that no "AssertationError" occurs.

<!-- grade_id: task_1, points: 0, grade: false, solution: true -->
```python
# I have read through the introduction about Tools, Grading and Setup
read_introduction = True

# I have a Google account and can access Google Colaboratory under my own 
# account
access_colaboratory = True

# I have a GitHub account
github_account = True
```

<!-- grade_id: task_1_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert read_introduction == True, "The variable read_introduction is not set to True"
assert access_colaboratory == True, "The variable access_colaboratory is not set to True"
assert github_account == True, "The variable github_account is not set to True"
```

2. Also update the ``github_username`` and ``idm_credentials`` accordingly. Keep the **'** around **'**username**'** and **'**ab12abcd**'**.

<!-- grade_id: task_2, points: 0, grade: false, solution: true -->
```python
# enter your github username
github_username = 'suchira_suresh'

# enter your FAU idM-credentials
idm_credentials = 'yh07ajaw'
```

<!-- grade_id: task_2_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert isinstance(github_username, str), "The variable github_username is not of type string"
assert isinstance(idm_credentials, str), "The variable idm_credentials is not of type string"
```

### **Submission**
After you finished completing the code above, you need to submit your assignment. Submitting the assignment entails syncing your copy of the notebook with your private GitHub assignment repository (``Applied Data Science in Medicine & Psychology/ads_assignment_1-<username>``):
1. Save your notebook, click ``File`` $\rightarrow$ ``Save``
2. From within Google Colaboratory, click ``File`` $\rightarrow$ ``save a copy in Github``
3. Select ``Applied Data Science in Medicine & Psychology/ads_assignment_1-<username>`` from the repository drop-down menu
4. Type ``ads_assignment_1.ipynb`` into the File path text box
5. Optionally, you may enter a description into the "Commit message" text box (like: "solved task 1" or "ready for grading")
6. Make sure that the "Include a link to Colaboratory" box is checked
7. Press "Ok"

A new tab or window should open up, displaying your submitted notebook within your GitHub repository.

### **Updating your submission**
Prior to the assignment deadline you can re-submit your assignment as often as you would like. This will update the copy of your assignment file in your assignment's repository. Update the assignment from the notebook that you saved in your Google Drive (access from your Google Colab account) or your own repository.